---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Introduction

```yaml
type: VideoExercise
key: 7101b90e13
xp: 50
```

`@projector_key`
bc6ea08480a45ca5244c0238607c961c

---

## Insert exercise title here

```yaml
type: NormalExercise
key: d6df1956fb
xp: 100
```

Context should have at leat a lot of words

`@instructions`
Try to produce something that make sens please

`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
